```python
%matplotlib inline
```

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt

from nose.tools import *

from mpl_toolkits.basemap import Basemap

import networkx as nx
# Write your imports here
```

# Working with Spatial Data and Networks Lab
## Preparing data, extracting features, visualization and modelling

In this lab we'll explore terrorism, as described in [this](https://www.kaggle.com/START-UMD/gtd) dataset.

The full codebook, describing all variables, their values, the data collection process, and many more details, is located [here](http://start.umd.edu/gtd/downloads/Codebook.pdf).

We'll perform our usual process:
* Reading and inspecting the data
* Validating the data
* Exploring the data
* Creating inferences and getting to conclusions

### Problem 1. Read the dataset (1 point)
The dataset is provided in the `data` folder. Read it into Python. You'll likely get an encoding error, so you can use `encoding = "cp1252"`.

Use the column `eventid` as the index, since this is a good unique identifier and it's used in the dataset (e.g. the column `related` contains the `eventid` of all related crimes).

You'll see the dataset has a lot of columns. For this lab, we aren't going to need all of them, only a few. Select these columns (and rename them as described in the tests below):
* eventid
* iyear, imonth, iday
* country_txt
* region_txt
* multiple (1 if it's a part of a multiple incident and 0 otherwise)
* latitude, longitude
* targtype1_txt (first target type)
* weaptype1_text (weapon type)
* gname (perpetrator group name)
* nkill (total number of fatalities)
* related (IDs of the related incidents)

You'll have a total of 12 columns (the eventid which will become an index is not considered a typical column in a dataframe) and just over 170 thousand records.

<!-- grade_id: cell-1c744200bd2b910d, points: 0, grade: false, solution: true -->
```python
terrorism_data = pd.read_csv("data/globalterrorismdb.csv", encoding="cp1252",index_col="eventid")
terrorism_data = terrorism_data[["iyear","imonth","iday","country_txt","region_txt","multiple",
                                 "latitude","longitude","targtype1_txt","weaptype1_txt","gname", "nkill","related"]]

terrorism_data.columns = ["year","month","day","country","region","multiple","latitude","longitude","target_type",
                         "weapon_type","group_name","fatalities","related"]
```

<!-- grade_id: cell-5f34e9684b9bb47b, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(terrorism_data)
assert_list_equal(terrorism_data.columns.tolist(), ["year", "month", "day", "country", "region", "multiple", "latitude", "longitude", "target_type", "weapon_type", "group_name", "fatalities", "related"])
```

### Problem 2. Convert columns (1 point)
Convert the "year", "month" and "day" columns into a single date. Call the new column "date".

Note that some months and/or days contain 0 as their values. This means that the exact month or day is not present or known. In these cases, write "not a time" (NaT) values in the "date" column. Do not remove those records!

Also, convert the "multiple" column into a Boolean one.

<!-- grade_id: cell-c415bca15e14b990, points: 0, grade: false, solution: true -->
```python
terrorism_data["date"] = pd.to_datetime(terrorism_data.year * 10000 + terrorism_data.month * 100 + terrorism_data.day,format = "%Y%m%d",errors = "coerce")
terrorism_data.multiple = terrorism_data.multiple.astype(np.bool)
terrorism_data = terrorism_data.drop(["year", "month", "day"], axis = 1)

```

<!-- grade_id: cell-8f8c1f0a4e1f54be, points: 1, grade: true, solution: false -->
```python
assert_greater(len(terrorism_data.date), 0)
assert_equal(terrorism_data.multiple.dtype, np.bool)
```

### Problem 3. Show how the number of attacks evolves with time (1 point)
Group all incidents by year. Create a line plot showing how the attacks evolve.

Create a function which accepts a dataframe and returns the years with minimum and maximum incidents.

Use this function to see the numbers for your dataset.

<!-- grade_id: cell-d95cdafff350daaf, points: 0, grade: false, solution: true -->
```python
# Note: This is not autograded
attacks_by_year = terrorism_data.groupby(terrorism_data.date.dt.year).date.count()
attacks_by_year.plot()
plt.title("Number of attacks evolves through time")
plt.xlabel("year")
plt.ylabel("Quantity Attacks")
plt.show()
```

<!-- grade_id: cell-3c2352bef75801fa, points: 0, grade: false, solution: true -->
```python

def get_max_and_min_incidents_years(dataframe):
    # Don't forget to use the dataframe parameter!
    min_incidents_year, max_incidents_year = dataframe.groupby(dataframe.date.dt.year).date.count().sort_values(ascending=True).index[0], dataframe.groupby(dataframe.date.dt.year).date.count().sort_values(ascending=False).index[0]
    return (min_incidents_year, max_incidents_year)
```

<!-- grade_id: cell-9fab4a50c075077f, points: 1, grade: true, solution: false -->
```python
min_year, max_year = get_max_and_min_incidents_years(terrorism_data)
assert_equal(min_year, 1971)
assert_equal(max_year, 2014)
```

### Problem 4. Filter terror attacks (1 point)
Get only recent attacks (from 1 Jan 2000 onwards, inclusive). Save them in the variable `recent_attacks`.

Create another variable which contains attacks before that. Call that `older_attacks`.

We'll compare these later.

<!-- grade_id: cell-1c3587b0a671b81e, points: 0, grade: false, solution: true -->
```python
recent_attacks = terrorism_data[terrorism_data.date.dt.year >= 2000]
older_attacks = terrorism_data[terrorism_data.date.dt.year < 2000]
```

<!-- grade_id: cell-1af72cf0b455864f, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(recent_attacks)
assert_is_not_none(older_attacks)
```

### Problem 5. Compare recent to older attacks (1 point)
Create a function which accepts a dataframe and returns a `Series`. The index should be a country name, and the values should be the number of  attacks. Sort the values from highest to lowest.

<!-- grade_id: cell-8e2899086cca6575, points: 0, grade: false, solution: true -->
```python
def get_num_attacks_by_country(dataframe):
    series = dataframe.groupby("country").size().nlargest(10)
    return series
```

<!-- grade_id: cell-c51b7c39a69134f9, points: 1, grade: true, solution: false -->
```python
recent_attacks_by_country = get_num_attacks_by_country(recent_attacks)
older_attacks_by_country = get_num_attacks_by_country(older_attacks)
print(recent_attacks_by_country)
print(older_attacks_by_country)
assert_equal(recent_attacks_by_country.shape, (10,))
assert_equal(older_attacks_by_country.shape, (10,))
```

```python
def show_graph(num_attacks, title):
    plt.bar(range(len(num_attacks)), num_attacks)
    plt.xticks(range(len(num_attacks)), num_attacks.index, rotation = "vertical")
    plt.title(title)
    plt.show()

show_graph(older_attacks_by_country, "Older attacks (before 2000)")
show_graph(recent_attacks_by_country, "Recent attacks (since 2000)")
```

We can see that these attacks definitely changed very much. First, there are many more since 2000, and second, the top targeted countries have changed a lot: only India remains in the top chart (execute the following cell if you don't believe me).

```python
set(recent_attacks_by_country.index).intersection(set(older_attacks_by_country.index))
```

There's a lot more to explore in the data, but let's focus on the spatial aspect.

### Problem 6. Create a basic map (2 points)
Plot all places where terrorist attacks have occurred.

The skeleton below almost does the job. What you need to do is: first, convert the geographical coordinates to horizontal and assign them to the `x` and `y` variables; second, create a scatter plot, using a red circle (`"o"`), with marker size 1.

<!-- grade_id: cell-949ee28085e3121b, points: 0, grade: false, solution: true -->
```python
plt.figure(figsize = (12, 10))
m = Basemap(projection = "merc", llcrnrlat = -73, llcrnrlon = -180, urcrnrlat = 80, urcrnrlon = 180)
x, y = m(terrorism_data.longitude.tolist(),terrorism_data.latitude.tolist())
m.plot(x,y,"o",markersize = 1,color = "red")
m.drawcoastlines()
m.drawcountries()
m.fillcontinents(color = "lightgreen", lake_color = "aqua")
m.drawmapboundary(fill_color = "aqua")
plt.show()
```

<!-- grade_id: cell-db1a0f33e3c947b5, points: 2, grade: true, solution: false -->
```python
assert_is_not_none(x)
assert_is_not_none(y)

```

### Problem 7. Emphasize incidents which are part of a group (0 points)
Redraw the map above. This time, filter out the attacks which are part of a series (this is contained in the column `multiple`). Show them on the map with an orange color.

<!-- grade_id: cell-c92a19d6d8714acc, points: 0, grade: false, solution: true -->
```python
# Note: This is not autograded
single_attacks = terrorism_data[terrorism_data.multiple == False]
multiple_attacks = terrorism_data[terrorism_data.multiple == True]
```

<!-- grade_id: cell-0140c2050b753b80, points: 0, grade: true, solution: false -->
```python
plt.figure(figsize = (12, 10))
m = Basemap(projection = "merc", llcrnrlat = -73, llcrnrlon = -180, urcrnrlat = 80, urcrnrlon = 180)

x, y = m(single_attacks.longitude.tolist(), single_attacks.latitude.tolist())
m.plot(x, y, "o", markersize = 1, color = "red")

x, y = m(multiple_attacks.longitude.tolist(), multiple_attacks.latitude.tolist())
m.plot(x, y, "o", markersize = 1, color = "orange")

m.drawcoastlines()
m.drawcountries()
m.fillcontinents(color = "lightgreen", lake_color = "aqua")
m.drawmapboundary(fill_color = "aqua")
plt.show()
```

You can see that a lot of the attacks are parts of series, possibly organized. But, remember that we have a point on the map for every single attack. If we see how many multiple attacks there are (as a fraction of the whole), we'll get about 13.4%:

```python
len(multiple_attacks) / len(terrorism_data)
```

The plot we made was a little misleading. We're displaying where attacks occur but we omitted the "density" of attacks: each location gets a point no matter whether there was one incident, or ten thousand in that point. 

We can try plotting the densities instead.

### Problem 8. Plot attack densities (1 point)
First, we need to filter out attacks with unknown (NaN) location. Do this and save it in the variable `attacks_with_locations`.

The following cell prints a 2D histogram (also called a heatmap) over the world map. Execute it, it should run just fine.

<!-- grade_id: cell-e249aca4f843b587, points: 0, grade: false, solution: true -->
```python
attacks_with_locations = terrorism_data[["longitude","latitude"]].dropna()
```

```python
def plot_attack_densities(attack_data, title = "Attack densities"):
    plt.figure(figsize = (15, 10))
    m = Basemap(projection = "merc", llcrnrlat = -73, llcrnrlon = -180, urcrnrlat = 80, urcrnrlon = 180)

    # Prepare histogram bins
    num_bins = 100
    lon_bins = np.linspace(-180, 180, num_bins + 1)
    lat_bins = np.linspace(-90, 90, num_bins + 1)

    # Create 2D histogram values
    density, x_breaks, y_breaks = np.histogram2d(
        attack_data.latitude,
        attack_data.longitude, 
        [lat_bins, lon_bins])

    # Create the basis of the histogram - the (x, y) value pairs
    # and map them to 2D distances
    lon_bins_2d, lat_bins_2d = np.meshgrid(lon_bins, lat_bins)
    x, y = m(lon_bins_2d, lat_bins_2d)

    m.drawcoastlines()
    m.drawcountries()
    m.pcolormesh(x, y, density)
    m.colorbar()
    plt.title(title)
    plt.show()
```

```python
plot_attack_densities(attacks_with_locations)
```

We can see various "hot points". The hottest one, as expected from the data above, is Iraq. Actually, it has so many attacks that it overshadows most of the data in our dataset.

Let's remove it and plot the data again. This is a way to estimate **importance** of something (although not a very scientific one :)).

<!-- grade_id: cell-4696a4495fae2524, points: 0, grade: false, solution: true -->
```python
attacks_without_iraq = attacks_with_locations[attacks_with_locations.country != "Iraq"]
```

<!-- grade_id: cell-dffe18739baef50b, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(attacks_with_locations)
assert_is_not_none(attacks_without_iraq)
```

```python
plot_attack_densities(attacks_without_iraq, "Attack densities (Iraq omitted)")
```

Well, the map "lit up". Notice that the colorbar now shows completely different values.

Based on these visualizations, we can conclude that the number of attacks in Iraq is really significant compared to all others.

We're also able to see more hot points now, and we may try to interpret them. However, that's not what we're going to do.

### Problem 9. Create a graph of connected incidents (1 + 1 points)
We already have incidents which are part of groups: the variable `multiple_attacks` holds them.

However, it turns out that some of these hold NaNs in the `related` column.

Create a new dataset which holds the `eventid` and `related` incidents of every attack in `multiple_attacks`.

Filter the dataset so that it contains only valid (i.e. non-empty or NaN) `related` values.

Store this dataset into `incident_relations`.

<!-- grade_id: cell-a2682fbc00a9314f, points: 0, grade: false, solution: true -->
```python
incident_relations = multiple_attacks[["related"]].dropna()
```

<!-- grade_id: cell-2760002c8cd2e6fb, points: 1, grade: true, solution: false -->
```python
assert_equal(incident_relations.shape, (22840, 1))
```

Get all relations by splitting the column first, like so:
```
relations = incident_relations.related.str.split(", ")
```

For each relation, if the id of the first and second node aren't the same, add the relation as an edge. An edge has the format `(id_1, id_2)`.

The code is not very hard to understand but it definitely IS tricky to write. So I'm giving it to you for free :).

```python
edges = set()

relations = incident_relations.related.str.split("\D+")
for index, related in relations.iteritems():
    current_index = int(index)
    for related_id in related:
        if related_id != "":
            current_related = int(related_id)
            # (id_1, id_2) and (id_2, id_1) should be considered the same
            # undirected edge but in the dataset, they are separate
            # so we need to avoid double-adding edges
            if current_index == current_related:
                continue
            elif current_index > current_related:
                current_index, current_related = current_related, current_index
            edges.add((current_index, current_related))
print(len(edges))
```

Create a graph, `attack_graph` and add all edges that you (or I :D) just calculated to it.

Beware: visualization will take a lot of time, so **don't do that** (or just comment it out before submitting). I'm only testing basic properties on the graph.

<!-- grade_id: cell-cc347a3a786ca561, points: 0, grade: false, solution: true -->
```python
attack_graph = nx.Graph()
attack_graph.add_edges_from(edges)
```

<!-- grade_id: cell-a20bcc1a8794b893, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(attack_graph)
```